# Denison CS181/DA210 Homework 5.a - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

#### Setup

Note that for these exercises, we'll use `mysocket.py`, provided with the book, and available to you now in `modules/` in this repository.

```python
import os
import os.path
import sys
import importlib

module_dir = "../../modules"
module_path = os.path.abspath(module_dir)
if not module_path in sys.path:
    sys.path.append(module_path)

import mysocket as sock
importlib.reload(sock)
```

Output:
```
<module 'mysocket' from 'C:\\Users\\Hudson Pham\\Documents\\school\\denison sophomore\\CS181\\cs181-f22\\modules\\mysocket.py'>
```

---

## Part A: Identifying Resources with URLs and URIs

Uniform Resource Identifiers (URIs) and Uniform Resource Locators (URLs) define a standard notation for specifying the files, data, and resources of the internet.  Note that URI is the broader term, so all URLs are URIs.

Using an explicit protocol scheme, host location, and resource path, URLs can be used to uniquely identify a resource at a specific location on the internet.  These components are summarized in the following table:

Item | Description
:----|:--------------
_protocol_ | The network stack layer above TCP; we'll use `http` and `https`
_location_ | The server/host machine within the internet
_port_     | The program used for connections; we usually use port 80 for `HTTP` web server programs and port 443 for HTTPS web server programs
_resource-path_ | Identifies a particular resource within the host/port endpoint; could also include a query string

The general form of a URL is given by the following (shown with extra spaces for readability):

_protocol_ : // _location_ [ : _port_ ] _resource-path_

**Q1:** Type the following URL in a web browser: http://datasystems.denison.edu:80/topnames.html.  What are the _protocol_, _location_, _port_, and _resource-path_ for this URL?

Hint: Be very careful about specifying where any forward slashes ( `/` ) belong!

Here is what we get after hitting the provided URL in this URL bar in browser:

Protocol: https

Location: datasystems.denison.edu

Port: 443

Resource-path: /topnames.html

**Q2:** Now, use a search engine to search for "Denison University".  What are the _protocol_, _location_, _port_, and _resource-path_ for the resulting URL? 

https://www.google.com/search?q=Denison+University&sxsrf=ALiCzsbODEA8OET98kGQm5qc2ivSglvjRA%3A1669783005866&source=hp&ei=3d2GY8O_MY-cptQP47mWkAE&iflsig=AJiK0e8AAAAAY4br7f8MCsfqPWqNODKAuuis_17ABbV9&ved=0ahUKEwiDoqe3itX7AhUPjokEHeOcBRIQ4dUDCAw&uact=5&oq=Denison+University&gs_lcp=Cgdnd3Mtd2l6EAMyBAguECcyBQgAEIAEMgUIABCABDIFCAAQgAQyBQgAEIAEMgUIABCABDIFCAAQgAQyCwgAEIAEELEDEIMBMgUIABCABDIFCAAQgAQ6BwgjEOoCECc6BAgjECc6BQgAEJECOgsILhCABBDHARDRAzoOCC4QgAQQsQMQxwEQ0QM6CAgAEIAEELEDOhEILhCDARDHARCxAxDRAxCABDoICAAQsQMQgwE6CwguEIAEEMcBEK8BOgsILhCABBCxAxCDAToICC4QgAQQ1AI6CwguEMcBENEDEJECOg4IABCABBCxAxCDARDJAzoICC4QgAQQsQM6BQguEIAEOg0ILhCxAxDHARDRAxAKOgsILhCvARDHARCABFDPAVjuFWCBGGgBcAB4AIABYogBlQuSAQIxOJgBAKABAbABCg&sclient=gws-wiz

Protocol: https

Location: www.google.com

Port: 443

Resource-path: /search?q=Denison+University&sxsrf=ALiCzsbODEA8OET98kGQm5qc2ivSglvjRA%3A1669783005866&source=hp&ei=3d2GY8O_MY-cptQP47mWkAE&iflsig=AJiK0e8AAAAAY4br7f8MCsfqPWqNODKAuuis_17ABbV9&ved=0ahUKEwiDoqe3itX7AhUPjokEHeOcBRIQ4dUDCAw&uact=5&oq=Denison+University&gs_lcp=Cgdnd3Mtd2l6EAMyBAguECcyBQgAEIAEMgUIABCABDIFCAAQgAQyBQgAEIAEMgUIABCABDIFCAAQgAQyCwgAEIAEELEDEIMBMgUIABCABDIFCAAQgAQ6BwgjEOoCECc6BAgjECc6BQgAEJECOgsILhCABBDHARDRAzoOCC4QgAQQsQMQxwEQ0QM6CAgAEIAEELEDOhEILhCDARDHARCxAxDRAxCABDoICAAQsQMQgwE6CwguEIAEEMcBEK8BOgsILhCABBCxAxCDAToICC4QgAQQ1AI6CwguEMcBENEDEJECOg4IABCABBCxAxCDARDJAzoICC4QgAQQsQM6BQguEIAEOg0ILhCxAxDHARDRAxAKOgsILhCvARDHARCABFDPAVjuFWCBGGgBcAB4AIABYogBlQuSAQIxOJgBAKABAbABCg&sclient=gws-wiz

---

## Part B: HTTP Definition

Web browsers are simply programs that request data (often HTML of web pages) from web servers, and display them to the user.  HTTP exists to enable these requests.

As discussed in class, HTTP is an application protocol, and is therefore built on TCP and the sockets interface.

0. The web server program is in an ``always ready'' state, waiting with an unresolved TCP socket endpoint, listening for requests for port 80 (for HTTP).  
1. A client (e.g., your web browser or this notebook) makes a TCP connection to the server endpoint, and a bidrectional communication is initiated.  
2. The client constructs an _HTTP request_.  
3. The request is sent:  
3a. The request is sent over the TCP socket connection to the server.  
3b. The server receives the request and processes it, constructing an _HTTP response_.  
4. The response is sent:  
4a. The response is sent over the TCP socket connection back to the client.  
4b. The client receives the response and processes it.  
5. Both the client and server close the TCP socket connection.

Note that steps 2-4 can happen just once or many times, depending on the HTTP request parameters.

A module, `mysocket`, is included with our textbook, and import above as `import mysocket as sock`.  It provides the following helper functions:

Function                                           | Description
---------------------------------------------------|-------------------------------------------------------------------
`makeConnection(host, port)`                       | Establish a TCP connection from the client machine to a server at the given machine `host` and listening at the given `port`. This returns the socket connection.  This corresponds to Step 1 of the client-side steps.
`sendString(conn, s)`                              | Given an established socket `conn`, take `s`, a string, and send it over the connection.  This corresponds to Step 3 of the client-side steps, where `s` would define all the characters making up a complete HTTP request.
`receiveTillClose(conn)`                            | This performs a socket `recv()` from the connection, consuming data until the server closes the connection.  This returns the complete HTTP response message. This corresponds to Step 4 of the client-side steps, and assumes that a connection close will define the end of the response message.

-----------------------------------------------------------------------------------------------------------------------

Let's now walk through the steps of communication:

**Step 1**

```python
connection = sock.makeConnection("httpbin.org", 80)
assert connection is not None
```

**Step 2**

```python
request_line = 'GET / HTTP/1.1\r\n'     # we've already seen this
host_line = 'Host: httpbin.org\r\n'     # required for HTTP 1.1
one_and_done = 'Connection: close\r\n'  # specifies whether to keep connection alive
empty_line = '\r\n'                     # we need this before the (optional) body

request_message = request_line + host_line + \
                  one_and_done + empty_line
                  
print(request_message)
```

Output:
```
GET / HTTP/1.1
Host: httpbin.org
Connection: close
```

**Step 3**

```python
sock.sendString(connection, request_message)
```

**Step 4**

```python
response = sock.receiveTillClose(connection)
```

**Step 5**

```python
connection.close()
```

We can view the first 250 characters of the response (lines are separated by `'\r\n'`):

```python
print(response[:250])
```

Output:
```
HTTP/1.1 200 OK
Date: Wed, 30 Nov 2022 19:43:25 GMT
Content-Type: text/html; charset=utf-8
Content-Length: 9593
Connection: close
Server: gunicorn/19.9.0
Access-Control-Allow-Origin: *
Access-Control-Allow-Credentials: true

<!DOCTYPE html>
```

---

## Part C: Practice with HTTP Requests

**Q3:** Suppose we wish to retrieve (GET) a file via HTTP (so port 80) from `datasystems.denison.edu`.  The resource path of the file is `/data/ind0.json`.  We wish to use version 1.1 of HTTP and to request that the connection be closed after a single request/reply exchange.  We will need a header line to satisfy the HTTP 1.1 requirement of a valid `Host` header.  Write a sequence of code to compose a valid HTTP request as a Python string, and assign the result to `message`.

<!-- grade_id: cell-cae4365673329fd8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
request_line = 'GET /data/ind0.json HTTP/1.1\r\n'     
host_line = 'Host: datasystems.denison.edu\r\n'     
one_and_done = 'Connection: close\r\n'  
empty_line = '\r\n'                     

message = request_line + host_line + \
                  one_and_done + empty_line

print(message)
print("--------------------")
```

Output:
```
GET /data/ind0.json HTTP/1.1
Host: datasystems.denison.edu
Connection: close


--------------------
```

<!-- grade_id: cell-c37e0233f773840c, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert type(message) == str
assert message[:3] == "GET"
assert message[4:4+len("/data/ind0.json")] == "/data/ind0.json"
assert "Host: datasystems.denison.edu" in message
assert message.count('\r\n') == 4
assert message[-4:] == '\r\n\r\n'
```

**Q4:** Write a sequence of code to establish a connection to the host `datasystems.denison.edu` at port 80, to send the string `message` from the previous problem to the host, receive the reply from the host until the server closes the connection, assigning the reply to `reply`, and close the connection.  Note: if the request is not completely correct, a network connection can wait forever for a reply that will never come.  So if you have difficulty here, double check your answer to the previous problem.

<!-- grade_id: cell-59b25ae0233eb2f6, points: 0, grade: false, solution: true -->
```python
# Step 1
# YOUR CODE HERE
# raise NotImplementedError()
connection = sock.makeConnection("datasystems.denison.edu", 80)

# Step 2
# YOUR CODE HERE
# raise NotImplementedError()
sock.sendString(connection, message)

# Step 3
# YOUR CODE HERE
# raise NotImplementedError()
reply = sock.receiveTillClose(connection)

# Step 4
# YOUR CODE HERE
# raise NotImplementedError()
connection.close()

# Step 5
# YOUR CODE HERE
# raise NotImplementedError()

print(reply)
```

Output:
```
HTTP/1.1 200 OK
Date: Wed, 30 Nov 2022 19:43:25 GMT
Server: Apache/2.4.6 (CentOS)
Last-Modified: Wed, 16 Dec 2020 23:45:42 GMT
ETag: "10d-5b69d7922d580"
Accept-Ranges: bytes
Content-Length: 269
Connection: close
Content-Type: application/json

{"FRA": {"2007": {"pop": 64.02, "gdp": 2657.21}, "2017": {"pop": 66.87, "gdp": 2586.29}}, "GBR": {"2007": {"pop": 61.32, "gdp": 3084.12}, "2017": {"pop": 66.06, "gdp": 2637.87}}, "USA": {"2007": {"pop": 301.23, "gdp": 14451.9}, "2017": {"pop": 325.15, "gdp": 19485.4}}}
```

<!-- grade_id: cell-a262284d7dd83250, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert type(reply) == str
assert "200 OK" in reply
assert "application/json" in reply
assert reply.endswith("19485.4}}}")
```

**Q5:** Suppose we want to generalize the scenario from the first exercise, where the two things that can change are the *host location* and the *resource path*.  For example, we might want to change the host to `httpbin.org` and the resource path to `/`, or many other combinations.  Write a function
```
    buildRequest(location, resource)
```    
that constructs and returns a Python string containing a valid HTTP GET request that incorporates the parameters `location` and `resource` into the request at the appropriate places, and includes the appropriate header lines (for the required `Host` and to request the server close the connection after the exchange).

Note: Your function should not actually _issue_ the request.

<!-- grade_id: cell-fdcdf3fee2f5371c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def buildRequest(location, resource):
    request_line = 'GET {} HTTP/1.1\r\n'.format(resource)     
    host_line = 'Host: {}\r\n'.format(location) 
    one_and_done = 'Connection: close\r\n'  
    empty_line = '\r\n'                     

    return request_line + host_line + one_and_done + empty_line

print(buildRequest("httpbin.org", "/get"))
print("---------------------")
```

Output:
```
GET /get HTTP/1.1
Host: httpbin.org
Connection: close


---------------------
```

<!-- grade_id: cell-81a9009c478f82fa, points: 2, grade: true, solution: false -->
```python
# Testing cell

r1 = buildRequest("datasystems.denison.edu", "/data/ind0.json")
assert r1[:3] == "GET"
assert r1[4:4+len("/data/ind0.json")] == "/data/ind0.json"
assert "Host: datasystems.denison.edu" in r1
assert r1.count('\r\n') == 4
assert r1[-4:] == '\r\n\r\n'

r2 = buildRequest("httpbin.org", "/get")
assert r2[:3] == "GET"
assert r2[4:4+len("/get")] == "/get"
assert "Host: httpbin.org" in r2
assert r2.count('\r\n') == 4
assert r2[-4:] == '\r\n\r\n'
```

**Q6:** Write a function
```
    issueRequest(location, resource)
```
that first constructs a valid HTTP GET request for `resource` at host `location`, as a Python string (using your function from the previous question), and then performs the  request-reply steps of making the connection, sending the string request, receiving a reply until the connection closes, and finally closing the client side of the connection.  The function should return the reply.

<!-- grade_id: cell-b0ae35dd066cdd61, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def issueRequest(location, resource):
    connection = sock.makeConnection(location, 80)
    sock.sendString(connection, buildRequest(location, resource))
    reply = sock.receiveTillClose(connection)
    connection.close()
    return reply

print(issueRequest("datasystems.denison.edu", "/basic.html"))
```

Output:
```
HTTP/1.1 200 OK
Date: Wed, 30 Nov 2022 19:43:25 GMT
Server: Apache/2.4.6 (CentOS)
Accept-Ranges: bytes
Content-Length: 496
Connection: close
Content-Type: text/html; charset=UTF-8

<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Data Systems Basic HTML Page</title>
  </head>
  <body>
    <h1>First Level Heading</h1>

    <p>Paragraph defined in <b>body</b>.

    <h2>Second Level Heading</h2>

    <a href="http://docs.python.org">Link</a> to Python documentation.
    </p>

    <ul>
      <li>Item 1
      <ol>
        <li>Item 1 nested</li>
        <li>Item 2 nested</li>
      </ol>
      </li>
      <li>Item 2</li>
      <li>Item 3</li>
    </ul>
  </body>
</html>
```

```python
# Debugging cell #1
resp1 = issueRequest("datasystems.denison.edu", "/basic.html")
print(resp1)
```

Output:
```
HTTP/1.1 200 OK
Date: Wed, 30 Nov 2022 19:43:25 GMT
Server: Apache/2.4.6 (CentOS)
Accept-Ranges: bytes
Content-Length: 496
Connection: close
Content-Type: text/html; charset=UTF-8

<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Data Systems Basic HTML Page</title>
  </head>
  <body>
    <h1>First Level Heading</h1>

    <p>Paragraph defined in <b>body</b>.

    <h2>Second Level Heading</h2>

    <a href="http://docs.python.org">Link</a> to Python documentation.
    </p>

    <ul>
      <li>Item 1
      <ol>
        <li>Item 1 nested</li>
        <li>Item 2 nested</li>
      </ol>
      </li>
      <li>Item 2</li>
      <li>Item 3</li>
    </ul>
  </body>
</html>
```

```python
# Debugging cell #2
resp2 = issueRequest("datasystems.denison.edu", "/data/ind0.json")
print(resp2)
```

Output:
```
HTTP/1.1 200 OK
Date: Wed, 30 Nov 2022 19:43:25 GMT
Server: Apache/2.4.6 (CentOS)
Last-Modified: Wed, 16 Dec 2020 23:45:42 GMT
ETag: "10d-5b69d7922d580"
Accept-Ranges: bytes
Content-Length: 269
Connection: close
Content-Type: application/json

{"FRA": {"2007": {"pop": 64.02, "gdp": 2657.21}, "2017": {"pop": 66.87, "gdp": 2586.29}}, "GBR": {"2007": {"pop": 61.32, "gdp": 3084.12}, "2017": {"pop": 66.06, "gdp": 2637.87}}, "USA": {"2007": {"pop": 301.23, "gdp": 14451.9}, "2017": {"pop": 325.15, "gdp": 19485.4}}}
```

```python
# Debugging cell #3
resp3 = issueRequest("httpbin.org", "/get")
print(resp3)
```

Output:
```
HTTP/1.1 200 OK
Date: Wed, 30 Nov 2022 19:43:25 GMT
Content-Type: application/json
Content-Length: 197
Connection: close
Server: gunicorn/19.9.0
Access-Control-Allow-Origin: *
Access-Control-Allow-Credentials: true

{
  "args": {}, 
  "headers": {
    "Host": "httpbin.org", 
    "X-Amzn-Trace-Id": "Root=1-6387b25d-460b6ae809b6e20a0e160847"
  }, 
  "origin": "140.141.4.63", 
  "url": "http://httpbin.org/get"
}
```

<!-- grade_id: cell-558109be645f32b3, points: 2, grade: true, solution: false -->
```python
# Testing cell
resp1 = issueRequest("datasystems.denison.edu", "/basic.html")
assert "200 OK" in resp1
assert "text/html" in resp1
assert resp1.endswith("</html>\n")

resp2 = issueRequest("datasystems.denison.edu", "/data/ind0.json")
assert "200 OK" in resp2
assert "application/json" in resp2
assert resp2.endswith("19485.4}}}")

resp3 = issueRequest("httpbin.org", "/get")
assert "200 OK" in resp3
assert "application/json" in resp3
assert resp3.endswith(""""url": "http://httpbin.org/get"\n}\n""")
```

---

## Part D: HTTP Response Messages

The next set of exercises are about parsing through the reply resulting from a request.  If we consider an HTTP reply, we can partition it into a status line, the set of headers, and the body.  The exercises ask for functions that, given a reply, and parse the reply and return each of these pieces.

**Q7:** Write a function
```
    parseStatus(reply)
```
that finds and returns a Python string consisting of only the status line of a reply.  The returned value should include the line-terminating `"\r\n"`.

<!-- grade_id: cell-db4474c167be79ae, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def parseStatus(reply):
    return reply.split("\r\n")[0] + "\r\n"

reply = issueRequest("datasystems.denison.edu", "/basic.html")
print(repr(parseStatus(reply)))
reply = issueRequest("datasystems.denison.edu", "/foobar.txt")
print(repr(parseStatus(reply)))
```

Output:
```
'HTTP/1.1 200 OK\r\n'
'HTTP/1.1 404 Not Found\r\n'
```

<!-- grade_id: cell-4e8d16f523ad3818, points: 2, grade: true, solution: false -->
```python
r1 = issueRequest("datasystems.denison.edu", "/basic.html")
s1 = parseStatus(r1)
assert s1 == "HTTP/1.1 200 OK\r\n"

r2 = issueRequest("datasystems.denison.edu", "/foobar.txt")
s2 = parseStatus(r2)
assert s2 == "HTTP/1.1 404 Not Found\r\n"
```

**Q8:** Write a function
```
    parseHeaders(reply)
```
that finds and returns a single Python string that starts with the first header in the reply and continues up through the last header in the reply, including the line-terminating `"\r\n"`, but *not* the empty line separating the headers from the body.

<!-- grade_id: cell-4985de7e00373e2d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def parseHeaders(reply):
    a = reply.split("\r\n")
    aa = a[:7]
    return "\r\n".join(aa)
    
reply = issueRequest("datasystems.denison.edu", "/basic.html")
print(repr(parseHeaders(reply)))
reply = issueRequest("datasystems.denison.edu", "/foobar.txt")
print(repr(parseHeaders(reply)))
```

Output:
```
'HTTP/1.1 200 OK\r\nDate: Wed, 30 Nov 2022 19:43:25 GMT\r\nServer: Apache/2.4.6 (CentOS)\r\nAccept-Ranges: bytes\r\nContent-Length: 496\r\nConnection: close\r\nContent-Type: text/html; charset=UTF-8'
'HTTP/1.1 404 Not Found\r\nDate: Wed, 30 Nov 2022 19:43:25 GMT\r\nServer: Apache/2.4.6 (CentOS)\r\nContent-Length: 296\r\nConnection: close\r\nContent-Type: text/html; charset=iso-8859-1\r\n'
```

<!-- grade_id: cell-c10750ab8472692f, points: 2, grade: true, solution: false -->
```python
# Testing cell

r1 = issueRequest("datasystems.denison.edu", "/basic.html")
h1 = parseHeaders(r1)
assert "Server: Apache" in h1
assert "Connection: close\r\n" in h1
assert "Content-Type: text/html" in h1

r2 = issueRequest("datasystems.denison.edu", "/foobar.txt")
h2 = parseHeaders(r2)
assert "Server: Apache" in h2
assert "Connection: close\r\n" in h2
assert "Content-Type: text/html" in h2
```

**Q9:** Write a function
```
    parseBody(reply)
```
that finds and returns a single Python string that starts with the beginning of the body (i.e. after the empty line of the reply) and continues to the end of the reply.

<!-- grade_id: cell-fd88ef8bb50a41c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def parseBody(reply):
    a = reply.split("\r\n")
    if "OK" in a[0]:
        return "\r\n".join(a[8:])
    return "\r\n".join(a[7:])

reply = issueRequest("datasystems.denison.edu", "/basic.html")
print(parseBody(reply))
reply = issueRequest("datasystems.denison.edu", "/foobar.txt")
print(parseBody(reply))
```

Output:
```
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Data Systems Basic HTML Page</title>
  </head>
  <body>
    <h1>First Level Heading</h1>

    <p>Paragraph defined in <b>body</b>.

    <h2>Second Level Heading</h2>

    <a href="http://docs.python.org">Link</a> to Python documentation.
    </p>

    <ul>
      <li>Item 1
      <ol>
        <li>Item 1 nested</li>
        <li>Item 2 nested</li>
      </ol>
      </li>
      <li>Item 2</li>
      <li>Item 3</li>
    </ul>
  </body>
</html>

<!DOCTYPE HTML PUBLIC "-//IETF//DTD HTML 2.0//EN">
<html><head>
<title>404 Not Found</title>
</head><body>
<h1>Not Found</h1>
<p>The requested URL /foobar.txt was not found on this server.</p>
<hr>
<address>Apache/2.4.6 (CentOS) Server at datasystems.denison.edu Port 80</address>
</body></html>
```

<!-- grade_id: cell-7859319f27b6d899, points: 2, grade: true, solution: false -->
```python
# Testing cell
r1 = issueRequest("datasystems.denison.edu", "/basic.html")
b1 = parseBody(r1)
r2 = issueRequest("datasystems.denison.edu", "/foobar.txt")
b2 = parseBody(r2)
assert b1.startswith("<!DOCTYPE html>")
assert b1.endswith("</html>\n")
assert b2.startswith("<!DOCTYPE HTML")
assert b2.endswith("</body></html>\n")
```

---

---

## Part E

**Q10:** How much time (in minutes/hours) did you spend on this homework assignment?

For about 40 minutes

**Q11:** Who was your partner for this assignment?  If you worked alone, say so instead.

I worked alone